<a href="https://colab.research.google.com/github/nneomacharity1/CE880-Lab-coursework-/blob/main/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

---
# Welcome to CE880
### This is your week-6 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
!pip install -U scikit-learn
import sys
import warnings

if not sys.warnoptions:
    warnings.simplefilter("ignore")
```

Output:
```
Looking in indexes: https://pypi.org/simple, https://us-python.pkg.dev/colab-wheels/public/simple/
Requirement already satisfied: scikit-learn in /usr/local/lib/python3.8/dist-packages (1.0.2)
Collecting scikit-learn
  Downloading scikit_learn-1.2.0-cp38-cp38-manylinux_2_17_x86_64.manylinux2014_x86_64.whl (9.7 MB)
[2K     [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m9.7/9.7 MB[0m [31m55.5 MB/s[0m eta [36m0:00:00[0m
[?25hRequirement already satisfied: scipy>=1.3.2 in /usr/local/lib/python3.8/dist-packages (from scikit-learn) (1.7.3)
Requirement already satisfied: threadpoolctl>=2.0.0 in /usr/local/lib/python3.8/dist-packages (from scikit-learn) (3.1.0)
Requirement already satisfied: joblib>=1.1.1 in /usr/local/lib/python3.8/dist-packages (from scikit-learn) (1.2.0)
Requirement already satisfied: numpy>=1.17.3 in /usr/local/lib/python3.8/dist-packages (from scikit-learn) (1.21.6)
Installing collected packages: scikit-learn
  Attempting uninstall: scikit-learn
    Found e
```

---
## Question 1: 

#### Student’s t-test
Tests whether the means of two independent samples are significantly different and return p-value. Please consider $\alpha = 0.05$ 

Assumptions

* Observations in each sample are independent and identically distributed (iid).
* Observations in each sample are normally distributed.
* Observations in each sample have the same variance.

Interpretation

* H0: the means of the samples are equal.
* H1: the means of the samples are unequal.

<!-- grade_id: my_ttest_ans, points: 0, grade: false, solution: true -->
```python
from scipy.stats import ttest_ind
data1 = [0.873, 2.817, 0.121, -0.945, -0.055, -1.436, 0.360, -1.478, -1.637, -1.869]
data2 = [1.142, -0.432, -0.938, -0.729, -0.846, -0.157, 0.500, 1.183, -1.075, -0.169]

def my_ttest(data1, data2):
    """Write a code to tests whether the means of two independent samples are significantly
    different and return p-value. Please consider 𝛼 = 0.05  """
    # YOUR CODE HERE
    t, p = ttest_ind(data1, data2)
    return p
    raise NotImplementedError()
```

<!-- grade_id: my_ttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_ttest(data1, data2), 0.748, rel_tol = 0.05)
```

---
## Question 2: 

#### Paired Student’s t-test
Tests whether the means of two paired samples are significantly different. Please consider $\alpha = 0.05$ 

Assumptions

* Observations in each sample are independent and identically distributed (iid).
* Observations in each sample are normally distributed.
* Observations in each sample have the same variance.
* Observations across each sample are paired.

Interpretation

* H0: the means of the samples are equal.
* H1: the means of the samples are unequal.

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
# Example of the Paired Student's t-test
from scipy.stats import ttest_rel
data1 = [0.873, 2.817, 0.121, -0.945, -0.055, -1.436, 0.360, -1.478, -1.637, -1.869]
data2 = [1.142, -0.432, -0.938, -0.729, -0.846, -0.157, 0.500, 1.183, -1.075, -0.169]


def my_pairedttest(data1, data2):
    """Write code to tests whether the means of two paired samples are 
    significantly different. Please consider 𝛼=0.05 """
    # YOUR CODE HERE
    t, p = ttest_rel(data1, data2)
    return p
    raise NotImplementedError()
```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_pairedttest(data1, data2), 0.746, rel_tol = 0.05)
```

---
## Question 3: 

#### Shapiro-Wilk Test
Tests whether a data sample has a Gaussian distribution.

Assumptions

* Observations in each sample are independent and identically distributed (iid).

Interpretation

* H0: the sample has a Gaussian distribution.
* H1: the sample does not have a Gaussian distribution.

<!-- grade_id: my_normalityttest_ans, points: 0, grade: false, solution: true -->
```python
from scipy.stats import shapiro
data = [0.873, 2.817, 0.121, -0.945, -0.055, -1.436, 0.360, -1.478, -1.637, -1.869]

def my_normalityttest(data):
    """Write code to to check the normality of the data. In other words, 
    you are checking if the data is normally distributed and return p value. 
    Hint: Please see what is shapiro test and how you can use it
    https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.shapiro.html"""
    # YOUR CODE HERE
    stat, p = shapiro(data)
    return p
    raise NotImplementedError()

my_normalityttest(data)
```

Output:
```
0.19340917468070984
```

<!-- grade_id: my_normalityttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_normalityttest(data), 0.193, rel_tol = 0.05)
```