<a href="https://colab.research.google.com/github/nneomacharity1/CE880-Lab-coursework-/blob/main/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

---
# Welcome to CE880
### This is your week-4 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

### Task: Complete all the three questions given as follows:

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
!pip install -U scikit-learn
```

Output:
```
Looking in indexes: https://pypi.org/simple, https://us-python.pkg.dev/colab-wheels/public/simple/
Requirement already satisfied: scikit-learn in /usr/local/lib/python3.8/dist-packages (1.0.2)
Collecting scikit-learn
  Downloading scikit_learn-1.2.0-cp38-cp38-manylinux_2_17_x86_64.manylinux2014_x86_64.whl (9.7 MB)
[2K     [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m9.7/9.7 MB[0m [31m30.8 MB/s[0m eta [36m0:00:00[0m
[?25hRequirement already satisfied: joblib>=1.1.1 in /usr/local/lib/python3.8/dist-packages (from scikit-learn) (1.2.0)
Requirement already satisfied: scipy>=1.3.2 in /usr/local/lib/python3.8/dist-packages (from scikit-learn) (1.7.3)
Requirement already satisfied: threadpoolctl>=2.0.0 in /usr/local/lib/python3.8/dist-packages (from scikit-learn) (3.1.0)
Requirement already satisfied: numpy>=1.17.3 in /usr/local/lib/python3.8/dist-packages (from scikit-learn) (1.21.6)
Installing collected packages: scikit-learn
  Attempting uninstall: scikit-learn
    Found e
```

```python
import pandas as pd
import numpy as np
from sklearn.cluster import DBSCAN
from sklearn.cluster import KMeans
from sklearn.datasets import make_circles, make_moons
from sklearn.metrics import adjusted_rand_score
from sklearn.preprocessing import StandardScaler
from sklearn.metrics import silhouette_score, adjusted_rand_score
from sklearn.decomposition import PCA
scaler = StandardScaler()
```

---
## Question 1: 
We have already loaded `moon` data for you using `make_moons` function. 

Write a code to perform `k-means` clustering on `scaled_features` variable and return `silhouette_score` as an output

<!-- grade_id: myScore_kmean_ans, points: 0, grade: false, solution: true -->
```python
# Let's make two moons, where n is the number of samples
features, true_labels = make_moons(n_samples=250, noise=0.05, random_state=42)
# fit_transform() method is basically the combination of fit method and transform method, 
# it is equivalent to fit().transform(). This method performs fit and transform on the 
# input data at a single time and converts the data points. 
# If we use fit and transform separate when we need both then it 
# will decrease the efficiency of the model so we use fit_transform() 
# which will do both the work.
scaled_features = scaler.fit_transform(features)
myscore = []

def myScore_kmean(scaled_features):
    """Write code to perform k-means clustering and print silhouette_score"""
    # YOUR CODE HERE
    # Define the number of clusters
    n_clusters = 3
    # Initialize the k-means model
    kmeans = KMeans(n_clusters=n_clusters)
    # Fit the model to the data
    kmeans.fit(scaled_features)
    # Predict the cluster labels
    labels = kmeans.predict(scaled_features)
    # Calculate the silhouette score
    myScore_Kmean = silhouette_score(scaled_features, labels)
    return myScore_Kmean
    raise NotImplementedError()
```

<!-- grade_id: myScore_kmean_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(myScore_kmean(scaled_features), 0.5, rel_tol=0.10)
```

Output:
```
/usr/local/lib/python3.8/dist-packages/sklearn/cluster/_kmeans.py:870: FutureWarning: The default value of `n_init` will change from 10 to 'auto' in 1.4. Set the value of `n_init` explicitly to suppress the warning
  warnings.warn(
```

---
## Question 2: 

Write a code to perform `DBSCAN` clustering algorithm on `scaled_features` and return `silhouette_score` as an output. Please set the `eps` = 0.2 and `min_samples` = 5

```python
# Let's make two moons, where n is the number of samples
from sklearn.datasets import make_moons
from sklearn.cluster import DBSCAN
from sklearn.preprocessing import StandardScaler
from sklearn.metrics import silhouette_score
import math
features, true_labels =  make_moons(n_samples=250, noise=.1)

def myScore_DBSCAN(scaled_features):
    """Write code to perform DBSCAN algorithm and print silhouette_score"""
    # YOUR CODE HERE
    
    #converting all to float
    scaled_feat = [[float(x) for x in row] for row in scaled_features]
    # Scaling the data
    scaler = StandardScaler()
    scaled_feat = scaler.fit_transform(features)

    # Performing the DBSCAN clustering
    clusterer = DBSCAN(eps=0.2, min_samples=5)
    clusters = clusterer.fit_predict(scaled_features)

    # Remove samples that are not assigned to a cluster
    scaled_features = [scaled_features[i] for i in range(len(scaled_features)) if clusters[i] != -1]
    clusters = [clusters[i] for i in range(len(clusters)) if clusters[i] != -1]
    
    # Compute the silhouette score
    silhouette = silhouette_score(scaled_features, clusters)
    print(silhouette)
    return silhouette
    raise NotImplementedError()
```

<!-- grade_id: myScore_DBSCAN_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
assert math.isclose(myScore_DBSCAN(scaled_features), 0.2, rel_tol=0.10)
```

Output:
```
0.19335289153528398
```

## Question 3: 

Complete this code perform `PCA`  on `iris data` and return `explained_variance_ratio_` as an output for the first two principal components (PCs)

<!-- grade_id: myPCA_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
from sklearn.decomposition import PCA


url = "https://archive.ics.uci.edu/ml/machine-learning-databases/iris/iris.data"
# load dataset into Pandas DataFrame
df = pd.read_csv(url, names=['sepal length','sepal width','petal length','petal width','target'])

features = ['sepal length', 'sepal width', 'petal length', 'petal width']
# Separating out the features
x = df.loc[:, features].values
# Separating out the target
y = df.loc[:,['target']].values
# Standardizing the features
x = StandardScaler().fit_transform(x)
    
def myPCA(x):
    """Complete this code return the explained variance ratio for the first two prinical 
    components after doing PCA on x """
    # YOUR CODE HERE
    pca = PCA(n_components=2)
    pcs = pca.fit_transform(x)
    explained_variance_ratio = pca.explained_variance_ratio_
    return explained_variance_ratio
    raise NotImplementedError()
```

<!-- grade_id: myPCA_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
x= myPCA(x)
y = [0.7277045209380135, 0.23030523267680633]
np.testing.assert_allclose(x, y, rtol=1e-5, atol=0)

```