<a href="https://colab.research.google.com/github/nneomacharity1/CE880-Lab-coursework-/blob/main/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

---
# Welcome to CE880
### This is your week-3 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

We have already loaded `iris` dataset for you. Now, write a function to calculate the `median` on the 1st column of the dataset. 
`Hint`: To find the median, you first need to sort the values in your sample. But, before that we have sliced, the first column of `iris` data into a variable name `sample` and data type is a list. Then pass `sample` list as a parameter in you `my_median` function.

The `median` of a sample of numeric data is the value that lies in the middle when we sort the data. The data may be sorted in ascending or descending order, the median remains the same. Please also consider what to do if there is no middle value.

Please don't use `np.median` function. Write your own custom function

<!-- grade_id: my_median_ans, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets
iris = datasets.load_iris()
sample = list(iris.data[:,0])

def my_median(sample):
    """Function for calculating median"""
    # YOUR CODE HERE
    sample.sort()
    # Calculate the median
    if len(sample) % 2 == 1:
    # If the length of the list is odd, return the middle value
       median = sample[len(sample) // 2]
    else:
    # If the length of the list is even, return the average of the two middle values
      median = (sample[len(sample) // 2 - 1] + sample[len(sample) // 2]) / 2
  
    return median
    raise NotImplementedError()
```

<!-- grade_id: my_median_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
assert my_median(sample) == 5.8
```

---

## Question 2:

`Please read me before attemping the question`

A quartile is a type of quantile. The first quartile `(Q1)`, is defined as the middle number between the smallest number and the median of the data set, the second quartile `(Q2)` – median of the given data set while the third quartile `(Q3)`, is the middle number between the median and the largest value of the data set.

Algorithm to find Quartiles :
Quartiles are calculated by the help of the median. If the number of entries is an even number i.e. of the form 2n, then, first quartile `(Q1)` is equal to the median of the `n` smallest entries and the third quartile `(Q3)` is equal to the median of the `n` largest entries.

If the number of entries is an odd number i.e. of the form `(2n + 1)`, then

* the first quartile (Q1) is equal to the median of the n smallest entries
* the third quartile (Q3) is equal to the median of the n largest entries
* the second quartile(Q2) is the same as the ordinary median.

Example: 
input = [75, 69, 56, 46, 47, 79, 92, 97, 89, 88, 36, 96, 105, 32, 116, 101, 79, 93, 91, 112]
After sorting the above data set:
[32, 36, 46, 47, 56, 69, 75, 79, 79, 88, 89, 91, 92, 93, 96, 97, 101, 105, 112, 116]
Here the total number of terms is 20.

The second quartile `(Q2)` or the median of the above data is `(88 + 89) / 2 = 88.5`
The first quartile `(Q1)` is median of first `n` i.e. 10 terms (or `n` i.e. 10 smallest values) = 62.5
The third quartile `(Q3)` is the median of `n` i.e. 10 largest values (or last `n` i.e. 10 values) = 96.5
Then, `IQR = Q3 – Q1 = 96.5 – 62.5 = 34.0`


Now, your job is to write a function to Caluate Interquartile Range on the second column of `iris` dataset, which you have loaded in the question 1.
`hint`: We have created a new variable `sample2` and sliced the second column of the `iris` data. Please explore `np.percentile` function

<!-- grade_id: my_quartile_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
sample2 = list(iris.data[:,1])
def my_quartile(sample2):
    """Caluate Interquartile Range of the given list."""
    # YOUR CODE HERE

  # Calculating the Interquartile Range of the second column
    q25, q75 = np.percentile(sample2, [25, 75])
    iqr = q75 - q25
    return iqr
    raise NotImplementedError()
```

<!-- grade_id: my_quartile_test, points: 2, grade: true, solution: false -->
```python
"""Check that correct answers here"""
import math
assert math.isclose(my_quartile(sample2), 0.5, rel_tol=0.05)
```

---
## Question 3: 

Write a function to calculate the `standard deviation` on the third column of `iris` dataset. We have created a new variable `sample3` and sliced the third column of the iris data. Please create a function that first calculate the variance then calculate `standard deviation`. Please do not use inbuilt `numpy.var` and `numpy.std` functions. Write your own codes. 

<!-- grade_id: my_std_ans, points: 0, grade: false, solution: true -->
```python
import math
sample3 = list(iris.data[:,2])
def my_std(sample3, ddof=0):
    """Write a program to calculate the standard deviation of 3rd column of iris dataset."""
    
    # YOUR CODE HERE
    # Calculating the mean of the third column
    mean = sum(sample3) / len(sample3)
    # Calculating the variance of the third column
    variance = sum((x - mean) ** 2 for x in sample3) / (len(sample3) - 1)
   # Calculating the standard deviation of the third column
    my_std = variance ** 0.5
    return my_std
    raise NotImplementedError()
```

<!-- grade_id: my_std_test, points: 1, grade: true, solution: false -->
```python
"""Test you solution"""
assert math.isclose(my_std(sample3), 1.759, rel_tol=0.05)
```

---
## Question 4:

Write a program to find the mean of the 4th column of `iris` dataset. 

<!-- grade_id: my_mean_ans, points: 0, grade: false, solution: true -->
```python
sample4 = list(iris.data[:,3])

def my_mean(sample4):
    """write a function to find the mean of the 4th column of iris data"""
    # YOUR CODE HERE
    my_mean = sum(sample4) / len(sample4)
    return my_mean
    raise NotImplementedError()
```

<!-- grade_id: my_mean_test, points: 1, grade: true, solution: false -->
```python
"""Test you solution"""
assert math.isclose(my_mean(sample4), 1.19, rel_tol=0.05)
```

## Question 5

Write a Pandas program to calculate the correlation between column 1 and column 3 of `iris` dataset. `Hint` use `from scipy import stats` to find a Pearson correlation coefficient. 

<!-- grade_id: iris_correlation_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
from scipy import stats
from scipy.stats import pearsonr
sample1 = list(iris.data[:,0])
sample2 = list(iris.data[:,2])
def iris_correlation(sample1, sample2):
    """calculate the correlation"""
    # YOUR CODE HERE
    #df = pd.DataFrame(["sample1"], ["sample2"])
    # Calculating the correlation between sample1 and sample2
    corr, p_value = pearsonr(sample1,sample2)
    print(corr)
    print(p_value)
    return corr
    raise NotImplementedError()
```

```python
"""Test you solution"""
assert math.isclose(iris_correlation(sample1, sample2), 0.87, rel_tol=0.05)
```

Output:
```
0.8717537758865831
1.0386674194498099e-47
```